# Assignment 1
For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web.



Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python

```

```python
import re
example_string = "Amy is 5 years old"
result = example_string.split(" ")
result
```

Output:
```
['Amy', 'is', '5', 'years', 'old']
```

```python

```

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    #raise NotImplementedError()
```

```python

```

```python
simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
Ruth and Peter, their parents, have 3 kids."""
res = re.findall("[A-Z]\w*",simple_string)
res
```

Output:
```
['Amy', 'Mary', 'Ruth', 'Peter']
```

```python

```

## Part A

Find a list of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""
    # YOUR CODE HERE
    #finda names by their uppercase 1st letter
    res = re.findall("[A-Z]\w*",simple_string)
    return res
    raise NotImplementedError()
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

```

```python

```

```python

```

```python
with open ("assets/grades.txt", "r") as file:
    grades = file.read()
grades
```

Output:
```
'Ronald Mayr: A\nBell Kassulke: B\nJacqueline Rupp: A \nAlexander Zeller: C\nValentina Denk: C \nSimon Loidl: B \nElias Jovanovic: B \nStefanie Weninger: A \nFabian Peer: C \nHakim Botros: B\nEmilie Lorentsen: B\nHerman Karlsen: C\nNathalie Delacruz: C\nCasey Hartman: C\nLily Walker : A\nGerard Wang: C\nTony Mcdowell: C\nJake Wood: B\nFatemeh Akhtar: B\nKim Weston: B\nNicholas Beatty: A\nKirsten Williams: C\nVaishali Surana: C\nCoby Mccormack: C\nYasmin Dar: B\nRomy Donnelly: A\nViswamitra Upandhye: B\nKendrick Hilpert: A\nKillian Kaufman: B\nElwood Page: B\nMukti Patel: A\nEmily Lesch: C\nElodie Booker: B\nJedd Kim: A\nAnnabel Davies: A\nAdnan Chen: B\nJonathan Berg: C\nHank Spinka: B\nAgnes Schneider: C\nKimberly Green: A\nLola-Rose Coates: C\nRose Christiansen: C\nShirley Hintz: C\nHannah Bayer: B'
```

```python
pattern = "(?:[\s]*)([\w ]*)(?::\ B)"
result = re.findall(pattern,grades)
result
```

Output:
```
['Bell Kassulke',
 'Simon Loidl',
 'Elias Jovanovic',
 'Hakim Botros',
 'Emilie Lorentsen',
 'Jake Wood',
 'Fatemeh Akhtar',
 'Kim Weston',
 'Yasmin Dar',
 'Viswamitra Upandhye',
 'Killian Kaufman',
 'Elwood Page',
 'Elodie Booker',
 'Adnan Chen',
 'Hank Spinka',
 'Hannah Bayer']
```

```python
len(result)
```

Output:
```
16
```

```python

```

## Part B

The dataset file in [assets/grades.txt](https://github.com/berserkhmdvhb/Data-Sciecne-with-Python/blob/main/Course1/Week1/Assignment/assets/grades.txt) contains a line separated list of people with their grade in 
a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re
def grades():
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()
    # YOUR CODE HERE
    ## ?: for excluding 1.Whitespaces(1st group) and 3.Students' B grade(3rd group)
    pattern = "(?:[\s]*)([\w ]*)(?::\ B)"
    result = re.findall(pattern,grades)
    return result
    raise NotImplementedError()
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16

```

```python

```

```python
with open("assets/logdata.txt", "r") as file:
    logdata = file.read()
logdata[0:4000]
```

Output:
```
'146.204.224.152 - feest6811 [21/Jun/2019:15:45:24 -0700] "POST /incentivize HTTP/1.1" 302 4622\n197.109.77.178 - kertzmann3129 [21/Jun/2019:15:45:25 -0700] "DELETE /virtual/solutions/target/web+services HTTP/2.0" 203 26554\n156.127.178.177 - okuneva5222 [21/Jun/2019:15:45:27 -0700] "DELETE /interactive/transparent/niches/revolutionize HTTP/1.1" 416 14701\n100.32.205.59 - ortiz8891 [21/Jun/2019:15:45:28 -0700] "PATCH /architectures HTTP/1.0" 204 6048\n168.95.156.240 - stark2413 [21/Jun/2019:15:45:31 -0700] "GET /engage HTTP/2.0" 201 9645\n71.172.239.195 - dooley1853 [21/Jun/2019:15:45:32 -0700] "PUT /cutting-edge HTTP/2.0" 406 24498\n180.95.121.94 - mohr6893 [21/Jun/2019:15:45:34 -0700] "PATCH /extensible/reinvent HTTP/1.1" 201 27330\n144.23.247.108 - auer7552 [21/Jun/2019:15:45:35 -0700] "POST /extensible/infrastructures/one-to-one/enterprise HTTP/1.1" 100 22921\n2.179.103.97 - lind8584 [21/Jun/2019:15:45:36 -0700] "POST /grow/front-end/e-commerce/robust HTTP/2.0" 304 14641\n241.114.1
```

```python
pattern = """
(?P<host>[n\d.]*)([- ]*)
(?P<user_name>[\w\d]*)
(\ \[)
(?P<time>[\d]*/[\w]*/[\d :-]*)
(\]\ \")
(?P<request>[\w\d\s/%+.-]*)
"""
L = []
for item in re.finditer(pattern,logdata,re.VERBOSE):
    L.append(item.groupdict())  
# We can get the dictionary returned for the item with .groupdict()
L[0:50]
```

Output:
```
[{'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'},
 {'host': '197.109.77.178',
  'user_name': 'kertzmann3129',
  'time': '21/Jun/2019:15:45:25 -0700',
  'request': 'DELETE /virtual/solutions/target/web+services HTTP/2.0'},
 {'host': '156.127.178.177',
  'user_name': 'okuneva5222',
  'time': '21/Jun/2019:15:45:27 -0700',
  'request': 'DELETE /interactive/transparent/niches/revolutionize HTTP/1.1'},
 {'host': '100.32.205.59',
  'user_name': 'ortiz8891',
  'time': '21/Jun/2019:15:45:28 -0700',
  'request': 'PATCH /architectures HTTP/1.0'},
 {'host': '168.95.156.240',
  'user_name': 'stark2413',
  'time': '21/Jun/2019:15:45:31 -0700',
  'request': 'GET /engage HTTP/2.0'},
 {'host': '71.172.239.195',
  'user_name': 'dooley1853',
  'time': '21/Jun/2019:15:45:32 -0700',
  'request': 'PUT /cutting-edge HTTP/2.0'},
 {'host': '180.95.121.94',
  'user_name': 'mohr6893',
  'time': '21/Jun/2019:15:45:34 -0700'
```

```python
len(L) == 979
```

Output:
```
True
```

```python
one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}

L[0] == one_item
```

Output:
```
True
```

```python

```

```python

```

```python

```

```python

```

## Part C

Consider the standard web log file in [assets/logdata.txt](https://github.com/berserkhmdvhb/Data-Sciecne-with-Python/blob/main/Course1/Week1/Assignment/assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:
* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re
def logs():
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
    # YOUR CODE HERE
    
    pattern = """
    (?P<host>[n\d.]*)([- ]*)
    (?P<user_name>[\w\d]*)
    (\ \[)
    (?P<time>[\d]*/[\w]*/[\d :-]*)
    (\]\ \")
    (?P<request>[\w\d\s/%+.-]*)
    """
    L = []
    for item in re.finditer(pattern,logdata,re.VERBOSE):
        L.append(item.groupdict())
    for l in L:
        if l['user_name']=='':
            l['user_name'] = '-'    
    # We can get the dictionary returned for the item with .groupdict()
    return L
    raise NotImplementedError()
```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

```

```python

```

```python

```

```python
logs()[0:50]
```

Output:
```
[{'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'},
 {'host': '197.109.77.178',
  'user_name': 'kertzmann3129',
  'time': '21/Jun/2019:15:45:25 -0700',
  'request': 'DELETE /virtual/solutions/target/web+services HTTP/2.0'},
 {'host': '156.127.178.177',
  'user_name': 'okuneva5222',
  'time': '21/Jun/2019:15:45:27 -0700',
  'request': 'DELETE /interactive/transparent/niches/revolutionize HTTP/1.1'},
 {'host': '100.32.205.59',
  'user_name': 'ortiz8891',
  'time': '21/Jun/2019:15:45:28 -0700',
  'request': 'PATCH /architectures HTTP/1.0'},
 {'host': '168.95.156.240',
  'user_name': 'stark2413',
  'time': '21/Jun/2019:15:45:31 -0700',
  'request': 'GET /engage HTTP/2.0'},
 {'host': '71.172.239.195',
  'user_name': 'dooley1853',
  'time': '21/Jun/2019:15:45:32 -0700',
  'request': 'PUT /cutting-edge HTTP/2.0'},
 {'host': '180.95.121.94',
  'user_name': 'mohr6893',
  'time': '21/Jun/2019:15:45:34 -0700'
```

```python
##Verify answer by solution length
```

```python
solution_length=[[15, 9, 26, 26],
 [14, 13, 26, 54],
 [15, 11, 26, 61],
 [13, 9, 26, 29],
 [14, 9, 26, 20],
 [14, 10, 26, 26],
 [13, 8, 26, 35],
 [14, 8, 26, 63],
 [12, 8, 26, 47],
 [15, 9, 26, 34],
 [12, 11, 26, 58],
 [12, 9, 26, 50],
 [15, 14, 26, 30],
 [15, 12, 26, 36],
 [12, 11, 26, 44],
 [13, 12, 26, 27],
 [13, 9, 26, 31],
 [14, 7, 26, 25],
 [13, 12, 26, 26],
 [14, 11, 26, 48],
 [15, 9, 26, 22],
 [13, 9, 26, 41],
 [13, 11, 26, 32],
 [12, 11, 26, 43],
 [14, 11, 26, 47],
 [11, 7, 26, 37],
 [14, 1, 26, 25],
 [13, 10, 26, 34],
 [15, 1, 26, 48],
 [14, 10, 26, 45],
 [14, 1, 26, 31],
 [12, 9, 26, 25],
 [14, 8, 26, 66],
 [15, 1, 26, 42],
 [13, 1, 26, 42],
 [14, 1, 26, 29],
 [10, 1, 26, 31],
 [12, 11, 26, 34],
 [15, 1, 26, 55],
 [12, 11, 26, 52],
 [11, 1, 26, 24],
 [15, 1, 26, 50],
 [14, 11, 26, 31],
 [13, 11, 26, 41],
 [12, 1, 26, 49],
 [13, 9, 26, 23],
 [13, 1, 26, 68],
 [13, 11, 26, 62],
 [13, 1, 26, 22],
 [13, 1, 26, 52],
 [15, 11, 26, 46],
 [14, 1, 26, 55],
 [14, 10, 26, 19],
 [15, 10, 26, 56],
 [14, 12, 26, 49],
 [14, 13, 26, 38],
 [12, 9, 26, 23],
 [15, 1, 26, 32],
 [13, 11, 26, 23],
 [15, 11, 26, 34],
 [13, 1, 26, 51],
 [14, 1, 26, 23],
 [13, 9, 26, 50],
 [13, 9, 26, 45],
 [12, 11, 26, 49],
 [14, 1, 26, 47],
 [14, 11, 26, 36],
 [13, 1, 26, 49],
 [13, 12, 26, 55],
 [14, 1, 26, 24],
 [15, 8, 26, 32],
 [15, 1, 26, 52],
 [13, 10, 26, 36],
 [13, 1, 26, 55],
 [13, 11, 26, 48],
 [15, 15, 26, 37],
 [12, 9, 26, 28],
 [12, 9, 26, 45],
 [14, 1, 26, 36],
 [13, 1, 26, 57],
 [12, 1, 26, 35],
 [10, 1, 26, 43],
 [14, 1, 26, 44],
 [13, 1, 26, 26],
 [14, 1, 26, 37],
 [13, 13, 26, 54],
 [11, 8, 26, 28],
 [15, 8, 26, 57],
 [13, 1, 26, 45],
 [13, 8, 26, 37],
 [12, 1, 26, 41],
 [11, 9, 26, 29],
 [13, 9, 26, 50],
 [15, 8, 26, 34],
 [14, 11, 26, 56],
 [14, 1, 26, 37],
 [13, 12, 26, 25],
 [13, 1, 26, 39],
 [12, 1, 26, 46],
 [15, 1, 26, 32],
 [13, 13, 26, 22],
 [15, 12, 26, 33],
 [14, 12, 26, 35],
 [13, 10, 26, 48],
 [15, 1, 26, 39],
 [12, 17, 26, 27],
 [13, 13, 26, 40],
 [15, 9, 26, 37],
 [15, 13, 26, 26],
 [13, 11, 26, 46],
 [13, 1, 26, 22],
 [14, 9, 26, 49],
 [13, 1, 26, 20],
 [14, 12, 26, 23],
 [13, 12, 26, 22],
 [11, 1, 26, 31],
 [14, 9, 26, 33],
 [13, 1, 26, 55],
 [11, 1, 26, 25],
 [12, 8, 26, 66],
 [14, 1, 26, 31],
 [14, 12, 26, 29],
 [15, 8, 26, 42],
 [14, 1, 26, 44],
 [11, 1, 26, 39],
 [12, 11, 26, 56],
 [13, 10, 26, 42],
 [15, 1, 26, 22],
 [11, 10, 26, 36],
 [13, 1, 26, 31],
 [14, 12, 26, 64],
 [15, 1, 26, 25],
 [12, 11, 26, 30],
 [14, 10, 26, 18],
 [13, 10, 26, 33],
 [13, 10, 26, 56],
 [13, 1, 26, 44],
 [14, 1, 26, 35],
 [12, 9, 26, 29],
 [15, 1, 26, 42],
 [13, 1, 26, 42],
 [14, 11, 26, 54],
 [15, 1, 26, 42],
 [14, 1, 26, 26],
 [14, 1, 26, 36],
 [12, 8, 26, 43],
 [13, 11, 26, 55],
 [13, 11, 26, 51],
 [13, 10, 26, 52],
 [14, 1, 26, 23],
 [14, 1, 26, 23],
 [14, 1, 26, 26],
 [14, 1, 26, 52],
 [14, 1, 26, 65],
 [13, 1, 26, 49],
 [12, 12, 26, 59],
 [14, 10, 26, 40],
 [12, 9, 26, 38],
 [14, 1, 26, 32],
 [13, 9, 26, 25],
 [12, 10, 26, 24],
 [11, 10, 26, 45],
 [13, 10, 26, 27],
 [14, 10, 26, 36],
 [13, 8, 26, 65],
 [13, 1, 26, 31],
 [14, 1, 26, 26],
 [13, 17, 26, 60],
 [14, 11, 26, 44],
 [12, 1, 26, 45],
 [14, 1, 26, 27],
 [14, 9, 26, 44],
 [12, 10, 26, 47],
 [14, 11, 26, 32],
 [12, 9, 26, 49],
 [14, 10, 26, 27],
 [13, 9, 26, 38],
 [15, 1, 26, 33],
 [14, 1, 26, 64],
 [13, 11, 26, 52],
 [14, 1, 26, 22],
 [13, 1, 26, 24],
 [15, 1, 26, 28],
 [12, 10, 26, 40],
 [14, 1, 26, 20],
 [14, 1, 26, 55],
 [15, 1, 26, 37],
 [15, 1, 26, 25],
 [14, 1, 26, 52],
 [13, 8, 26, 68],
 [15, 1, 26, 23],
 [14, 10, 26, 24],
 [14, 9, 26, 27],
 [14, 9, 26, 35],
 [13, 13, 26, 42],
 [14, 11, 26, 42],
 [12, 1, 26, 50],
 [14, 1, 26, 35],
 [15, 8, 26, 38],
 [13, 10, 26, 40],
 [12, 1, 26, 23],
 [14, 9, 26, 33],
 [14, 9, 26, 53],
 [14, 9, 26, 43],
 [14, 1, 26, 47],
 [14, 1, 26, 25],
 [13, 11, 26, 18],
 [12, 9, 26, 42],
 [13, 1, 26, 39],
 [15, 1, 26, 33],
 [12, 12, 26, 55],
 [15, 11, 26, 50],
 [13, 11, 26, 49],
 [12, 1, 26, 32],
 [14, 1, 26, 56],
 [11, 1, 26, 34],
 [14, 11, 26, 50],
 [13, 1, 26, 29],
 [14, 10, 26, 37],
 [10, 1, 26, 33],
 [11, 8, 26, 17],
 [12, 10, 26, 50],
 [12, 1, 26, 23],
 [14, 12, 26, 28],
 [14, 11, 26, 40],
 [14, 9, 26, 30],
 [14, 1, 26, 37],
 [13, 1, 26, 52],
 [15, 9, 26, 48],
 [13, 14, 26, 57],
 [14, 1, 26, 36],
 [15, 1, 26, 41],
 [13, 1, 26, 31],
 [15, 1, 26, 20],
 [14, 10, 26, 47],
 [13, 1, 26, 47],
 [13, 11, 26, 41],
 [13, 1, 26, 54],
 [14, 9, 26, 54],
 [13, 10, 26, 23],
 [14, 1, 26, 42],
 [15, 10, 26, 21],
 [14, 11, 26, 47],
 [12, 1, 26, 43],
 [14, 13, 26, 26],
 [12, 1, 26, 34],
 [12, 11, 26, 40],
 [15, 9, 26, 51],
 [12, 15, 26, 23],
 [12, 1, 26, 31],
 [14, 10, 26, 23],
 [14, 1, 26, 42],
 [13, 1, 26, 33],
 [13, 1, 26, 33],
 [12, 14, 26, 45],
 [13, 1, 26, 36],
 [12, 1, 26, 41],
 [13, 9, 26, 58],
 [15, 1, 26, 22],
 [13, 10, 26, 25],
 [14, 1, 26, 45],
 [14, 1, 26, 43],
 [13, 1, 26, 20],
 [14, 1, 26, 35],
 [12, 1, 26, 47],
 [11, 9, 26, 32],
 [13, 1, 26, 51],
 [14, 1, 26, 35],
 [15, 9, 26, 28],
 [11, 10, 26, 53],
 [12, 10, 26, 47],
 [13, 13, 26, 20],
 [14, 12, 26, 24],
 [13, 1, 26, 45],
 [15, 9, 26, 43],
 [15, 1, 26, 20],
 [14, 1, 26, 49],
 [13, 9, 26, 22],
 [12, 10, 26, 24],
 [13, 1, 26, 47],
 [13, 8, 26, 42],
 [14, 1, 26, 50],
 [13, 10, 26, 48],
 [13, 1, 26, 40],
 [14, 1, 26, 40],
 [13, 1, 26, 56],
 [12, 11, 26, 51],
 [13, 8, 26, 36],
 [15, 10, 26, 24],
 [14, 12, 26, 35],
 [12, 1, 26, 24],
 [12, 1, 26, 30],
 [12, 1, 26, 58],
 [15, 1, 26, 62],
 [13, 1, 26, 34],
 [13, 1, 26, 26],
 [14, 9, 26, 50],
 [12, 1, 26, 63],
 [13, 12, 26, 34],
 [15, 8, 26, 47],
 [14, 10, 26, 24],
 [13, 1, 26, 23],
 [14, 1, 26, 21],
 [12, 12, 26, 60],
 [15, 1, 26, 36],
 [14, 1, 26, 51],
 [14, 1, 26, 32],
 [12, 13, 26, 23],
 [13, 8, 26, 32],
 [12, 15, 26, 50],
 [13, 1, 26, 37],
 [11, 8, 26, 26],
 [14, 1, 26, 37],
 [12, 10, 26, 24],
 [15, 1, 26, 54],
 [13, 1, 26, 19],
 [12, 1, 26, 35],
 [13, 1, 26, 54],
 [11, 10, 26, 30],
 [13, 8, 26, 34],
 [11, 1, 26, 54],
 [12, 9, 26, 55],
 [14, 9, 26, 37],
 [15, 12, 26, 25],
 [15, 11, 26, 49],
 [15, 13, 26, 24],
 [13, 13, 26, 42],
 [15, 1, 26, 36],
 [12, 1, 26, 27],
 [13, 10, 26, 21],
 [14, 12, 26, 32],
 [13, 1, 26, 29],
 [14, 1, 26, 42],
 [13, 1, 26, 42],
 [14, 1, 26, 54],
 [13, 15, 26, 46],
 [14, 10, 26, 24],
 [12, 1, 26, 24],
 [13, 9, 26, 26],
 [13, 10, 26, 41],
 [12, 9, 26, 24],
 [14, 13, 26, 32],
 [15, 1, 26, 33],
 [13, 13, 26, 35],
 [13, 1, 26, 48],
 [15, 9, 26, 32],
 [12, 11, 26, 62],
 [14, 10, 26, 45],
 [11, 1, 26, 53],
 [14, 1, 26, 24],
 [14, 1, 26, 58],
 [14, 11, 26, 43],
 [14, 1, 26, 22],
 [12, 10, 26, 29],
 [14, 7, 26, 22],
 [12, 1, 26, 30],
 [14, 1, 26, 49],
 [12, 1, 26, 58],
 [13, 1, 26, 30],
 [13, 11, 26, 23],
 [15, 12, 26, 25],
 [15, 8, 26, 47],
 [13, 8, 26, 24],
 [13, 13, 26, 65],
 [13, 1, 26, 54],
 [14, 9, 26, 41],
 [15, 9, 26, 40],
 [15, 10, 26, 42],
 [13, 1, 26, 44],
 [13, 1, 26, 43],
 [12, 12, 26, 36],
 [12, 9, 26, 30],
 [14, 1, 26, 43],
 [14, 12, 26, 20],
 [13, 1, 26, 54],
 [14, 8, 26, 47],
 [9, 1, 26, 37],
 [15, 1, 26, 59],
 [13, 1, 26, 19],
 [14, 1, 26, 28],
 [14, 1, 26, 31],
 [14, 1, 26, 22],
 [15, 1, 26, 37],
 [13, 10, 26, 35],
 [14, 1, 26, 40],
 [13, 1, 26, 40],
 [15, 13, 26, 38],
 [13, 1, 26, 39],
 [13, 11, 26, 28],
 [14, 1, 26, 29],
 [11, 1, 26, 59],
 [12, 10, 26, 43],
 [15, 8, 26, 61],
 [12, 1, 26, 40],
 [14, 1, 26, 37],
 [12, 1, 26, 22],
 [13, 10, 26, 20],
 [14, 10, 26, 23],
 [13, 14, 26, 27],
 [14, 11, 26, 40],
 [13, 1, 26, 23],
 [12, 9, 26, 51],
 [15, 1, 26, 23],
 [14, 11, 26, 22],
 [12, 9, 26, 30],
 [13, 1, 26, 24],
 [13, 1, 26, 29],
 [10, 12, 26, 47],
 [14, 16, 26, 54],
 [15, 11, 26, 21],
 [11, 1, 26, 39],
 [14, 1, 26, 55],
 [14, 13, 26, 49],
 [14, 1, 26, 25],
 [14, 11, 26, 29],
 [14, 8, 26, 49],
 [13, 7, 26, 23],
 [14, 14, 26, 37],
 [14, 1, 26, 39],
 [15, 12, 26, 24],
 [15, 9, 26, 22],
 [12, 1, 26, 49],
 [15, 1, 26, 34],
 [12, 1, 26, 23],
 [13, 1, 26, 53],
 [13, 9, 26, 43],
 [13, 1, 26, 60],
 [13, 1, 26, 40],
 [13, 1, 26, 32],
 [11, 1, 26, 51],
 [13, 9, 26, 43],
 [15, 14, 26, 54],
 [15, 10, 26, 39],
 [14, 11, 26, 33],
 [14, 11, 26, 40],
 [13, 1, 26, 31],
 [13, 9, 26, 55],
 [11, 11, 26, 22],
 [13, 11, 26, 45],
 [13, 1, 26, 51],
 [13, 9, 26, 27],
 [13, 1, 26, 51],
 [12, 1, 26, 46],
 [12, 1, 26, 42],
 [15, 12, 26, 51],
 [14, 1, 26, 48],
 [13, 1, 26, 46],
 [14, 1, 26, 39],
 [12, 13, 26, 23],
 [14, 1, 26, 46],
 [11, 1, 26, 24],
 [12, 8, 26, 50],
 [15, 11, 26, 56],
 [15, 9, 26, 23],
 [14, 12, 26, 41],
 [13, 1, 26, 27],
 [11, 1, 26, 34],
 [13, 1, 26, 44],
 [13, 1, 26, 27],
 [14, 10, 26, 25],
 [14, 11, 26, 26],
 [12, 1, 26, 48],
 [12, 1, 26, 55],
 [12, 1, 26, 44],
 [9, 1, 26, 31],
 [14, 12, 26, 36],
 [15, 1, 26, 26],
 [14, 9, 26, 32],
 [12, 11, 26, 24],
 [10, 1, 26, 32],
 [13, 12, 26, 38],
 [13, 1, 26, 58],
 [14, 1, 26, 22],
 [13, 10, 26, 38],
 [14, 1, 26, 26],
 [15, 11, 26, 48],
 [12, 1, 26, 24],
 [14, 10, 26, 23],
 [13, 8, 26, 45],
 [13, 11, 26, 51],
 [11, 1, 26, 30],
 [14, 1, 26, 49],
 [14, 9, 26, 59],
 [11, 1, 26, 25],
 [13, 1, 26, 39],
 [15, 11, 26, 56],
 [14, 8, 26, 25],
 [13, 1, 26, 21],
 [12, 1, 26, 22],
 [14, 1, 26, 30],
 [12, 11, 26, 43],
 [13, 11, 26, 33],
 [14, 1, 26, 56],
 [13, 11, 26, 47],
 [15, 1, 26, 48],
 [14, 1, 26, 63],
 [12, 1, 26, 48],
 [13, 10, 26, 38],
 [13, 8, 26, 45],
 [14, 9, 26, 29],
  [13, 10, 26, 23],
 [14, 9, 26, 24],
 [15, 1, 26, 53],
 [12, 1, 26, 33],
 [14, 1, 26, 48],
 [13, 10, 26, 31],
 [12, 13, 26, 46],
 [14, 8, 26, 59],
 [15, 12, 26, 23],
 [14, 10, 26, 49],
 [13, 1, 26, 33],
 [11, 1, 26, 35],
 [12, 11, 26, 22],
 [12, 1, 26, 22],
 [14, 11, 26, 42],
 [13, 1, 26, 44],
 [15, 1, 26, 37],
 [15, 12, 26, 35],
 [14, 1, 26, 47],
 [14, 13, 26, 40],
 [14, 13, 26, 52],
 [12, 10, 26, 29],
 [13, 11, 26, 49],
 [13, 9, 26, 53],
 [13, 1, 26, 54],
 [13, 1, 26, 56],
 [14, 1, 26, 49],
 [14, 11, 26, 31],
 [13, 1, 26, 47],
 [14, 1, 26, 54],
 [14, 11, 26, 26],
 [15, 1, 26, 55],
 [12, 1, 26, 34],
 [12, 1, 26, 32],
 [14, 11, 26, 24],
 [14, 1, 26, 33],
 [14, 12, 26, 32],
 [11, 1, 26, 44],
 [13, 10, 26, 25],
 [13, 1, 26, 39],
 [12, 14, 26, 52],
 [14, 9, 26, 26],
 [13, 1, 26, 25],
 [15, 1, 26, 54],
 [13, 1, 26, 32],
 [14, 9, 26, 18],
 [14, 10, 26, 54],
 [13, 1, 26, 57],
 [12, 10, 26, 47],
 [13, 1, 26, 53],
 [13, 1, 26, 25],
 [15, 10, 26, 21],
 [14, 1, 26, 60],
 [14, 1, 26, 23],
 [14, 1, 26, 60],
 [14, 11, 26, 26],
 [14, 14, 26, 26],
 [13, 1, 26, 24],
 [11, 1, 26, 56],
 [12, 1, 26, 67],
 [14, 1, 26, 33],
 [14, 15, 26, 22],
 [13, 1, 26, 57],
 [13, 11, 26, 36],
 [13, 1, 26, 54],
 [13, 1, 26, 36],
 [13, 11, 26, 29],
 [15, 1, 26, 52],
 [13, 8, 26, 41],
 [13, 1, 26, 29],
 [11, 11, 26, 40],
 [13, 1, 26, 34],
 [12, 8, 26, 63],
 [13, 1, 26, 27],
 [14, 1, 26, 32],
 [13, 1, 26, 36],
 [14, 13, 26, 47],
 [14, 1, 26, 46],
 [11, 1, 26, 44],
 [14, 1, 26, 55],
 [14, 10, 26, 42],
 [13, 1, 26, 55],
 [13, 9, 26, 63],
 [15, 1, 26, 54],
 [14, 8, 26, 57],
 [11, 1, 26, 53],
 [15, 9, 26, 41],
 [14, 13, 26, 53],
 [14, 10, 26, 21],
 [14, 9, 26, 36],
 [14, 1, 26, 57],
 [14, 11, 26, 29],
 [13, 8, 26, 58],
 [13, 11, 26, 39],
 [14, 12, 26, 38],
 [13, 1, 26, 25],
 [14, 11, 26, 50],
 [14, 1, 26, 48],
 [11, 1, 26, 36],
 [12, 1, 26, 38],
 [15, 1, 26, 51],
 [12, 1, 26, 54],
 [14, 8, 26, 28],
 [13, 11, 26, 44],
 [13, 1, 26, 49],
 [13, 1, 26, 65],
 [11, 1, 26, 38],
 [13, 1, 26, 25],
 [15, 1, 26, 50],
 [13, 1, 26, 46],
 [14, 1, 26, 63],
 [14, 9, 26, 30],
 [14, 10, 26, 23],
 [13, 9, 26, 32],
 [14, 1, 26, 21],
 [13, 1, 26, 31],
 [14, 10, 26, 18],
 [15, 1, 26, 41],
 [14, 1, 26, 51],
 [11, 1, 26, 52],
 [15, 1, 26, 33],
 [12, 9, 26, 26],
 [15, 12, 26, 24],
 [11, 1, 26, 47],
 [13, 13, 26, 39],
 [15, 1, 26, 31],
 [11, 9, 26, 36],
 [13, 12, 26, 38],
 [11, 9, 26, 45],
 [13, 1, 26, 35],
 [14, 8, 26, 44],
 [13, 1, 26, 53],
 [13, 1, 26, 24],
 [12, 1, 26, 47],
 [13, 9, 26, 31],
 [15, 1, 26, 24],
 [13, 10, 26, 22],
 [14, 1, 26, 49],
 [12, 1, 26, 66],
 [13, 9, 26, 21],
 [14, 11, 26, 38],
 [13, 10, 26, 46],
 [14, 1, 26, 52],
 [13, 1, 26, 21],
 [13, 1, 26, 29],
 [12, 11, 26, 25],
 [14, 10, 26, 47],
 [13, 1, 26, 50],
 [14, 9, 26, 53],
 [13, 9, 26, 35],
 [13, 11, 26, 34],
 [14, 1, 26, 38],
 [12, 10, 26, 58],
 [14, 1, 26, 70],
 [13, 1, 26, 54],
 [14, 10, 26, 46],
 [15, 1, 26, 32],
 [13, 1, 26, 34],
 [13, 9, 26, 39],
 [13, 1, 26, 35],
 [14, 9, 26, 52],
 [14, 11, 26, 47],
 [11, 9, 26, 49],
 [14, 10, 26, 24],
 [15, 9, 26, 30],
 [13, 1, 26, 68],
 [14, 9, 26, 28],
 [14, 11, 26, 43],
 [14, 1, 26, 57],
 [14, 1, 26, 44],
 [12, 12, 26, 40],
 [14, 1, 26, 44],
 [11, 11, 26, 38],
 [13, 1, 26, 42],
 [13, 1, 26, 46],
 [14, 1, 26, 24],
 [14, 11, 26, 20],
 [14, 1, 26, 31],
 [11, 9, 26, 43],
 [13, 10, 26, 30],
 [15, 11, 26, 36],
 [14, 1, 26, 41],
 [12, 1, 26, 23],
 [12, 11, 26, 32],
 [14, 10, 26, 22],
 [14, 1, 26, 53],
 [14, 10, 26, 23],
 [15, 11, 26, 46],
 [12, 8, 26, 35],
 [14, 1, 26, 53],
 [13, 10, 26, 40],
 [12, 1, 26, 41],
 [14, 1, 26, 21],
 [12, 1, 26, 25],
 [13, 11, 26, 21],
 [13, 1, 26, 38],
 [13, 1, 26, 45],
 [14, 1, 26, 24],
 [15, 10, 26, 30],
 [13, 1, 26, 40],
 [14, 11, 26, 41],
 [15, 9, 26, 59],
 [14, 10, 26, 51],
 [14, 9, 26, 52],
 [12, 9, 26, 48],
 [13, 11, 26, 27],
 [13, 11, 26, 57],
 [13, 12, 26, 22],
 [14, 1, 26, 26],
 [12, 9, 26, 33],
 [15, 10, 26, 38],
 [15, 9, 26, 35],
 [13, 8, 26, 23],
 [14, 12, 26, 48],
 [13, 12, 26, 34],
 [12, 1, 26, 31],
 [13, 1, 26, 19],
 [14, 11, 26, 52],
 [15, 10, 26, 46],
 [13, 1, 26, 45],
 [15, 1, 26, 49],
 [13, 1, 26, 24],
 [13, 13, 26, 43],
 [15, 1, 26, 48],
 [15, 15, 26, 48],
 [14, 1, 26, 32],
 [12, 10, 26, 61],
 [13, 1, 26, 62],
 [14, 16, 26, 37],
 [12, 1, 26, 55],
 [13, 1, 26, 46],
 [14, 1, 26, 19],
 [13, 13, 26, 46],
 [13, 1, 26, 51],
 [15, 1, 26, 29],
 [14, 9, 26, 35],
 [12, 1, 26, 59],
 [13, 1, 26, 37],
 [13, 12, 26, 49],
 [14, 11, 26, 22],
 [12, 9, 26, 23],
 [11, 1, 26, 54],
 [15, 1, 26, 24],
 [12, 1, 26, 44],
 [13, 1, 26, 30],
 [12, 1, 26, 57],
 [12, 11, 26, 49],
 [14, 1, 26, 51],
 [12, 11, 26, 56],
 [14, 11, 26, 46],
 [15, 10, 26, 23],
 [14, 1, 26, 36],
 [12, 9, 26, 57],
 [14, 11, 26, 34],
 [15, 8, 26, 49],
 [14, 1, 26, 36],
 [13, 1, 26, 43],
 [12, 12, 26, 52],
 [14, 1, 26, 36],
 [13, 1, 26, 41],
 [13, 1, 26, 40],
 [14, 1, 26, 26],
 [13, 1, 26, 55],
 [14, 9, 26, 28],
 [14, 10, 26, 29],
 [13, 1, 26, 43],
 [14, 9, 26, 32],
 [14, 1, 26, 23],
 [15, 1, 26, 23],
 [13, 1, 26, 40],
 [11, 1, 26, 24],
 [14, 1, 26, 46],
 [14, 1, 26, 31],
 [14, 10, 26, 39],
 [15, 1, 26, 23],
 [12, 1, 26, 36],
 [14, 1, 26, 21],
 [14, 1, 26, 51],
 [15, 8, 26, 53],
 [14, 1, 26, 43],
 [13, 1, 26, 29],
 [13, 10, 26, 40],
 [13, 11, 26, 44],
 [14, 13, 26, 38],
 [12, 9, 26, 34],
 [15, 1, 26, 26],
 [13, 12, 26, 49],
 [14, 1, 26, 31],
 [14, 11, 26, 29],
 [14, 1, 26, 22],
 [13, 1, 26, 45],
 [13, 12, 26, 37],
 [13, 1, 26, 46],
 [13, 1, 26, 44],
 [12, 10, 26, 28],
 [13, 9, 26, 34],
 [15, 8, 26, 23],
 [14, 9, 26, 29],
 [12, 13, 26, 30],
 [12, 10, 26, 46],
 [13, 1, 26, 54],
 [14, 1, 26, 22],
 [11, 1, 26, 37],
 [14, 11, 26, 43],
 [14, 9, 26, 52],
 [13, 1, 26, 42],
 [13, 1, 26, 42],
 [13, 1, 26, 22],
 [14, 1, 26, 67],
 [13, 1, 26, 23],
 [14, 1, 26, 38],
 [14, 1, 26, 50],
 [13, 1, 26, 45],
 [13, 8, 26, 45],
 [15, 1, 26, 42],
 [15, 1, 26, 47],
 [15, 9, 26, 34],
 [14, 9, 26, 45],
 [15, 11, 26, 58],
 [13, 12, 26, 31],
 [12, 1, 26, 52],
 [14, 1, 26, 23],
 [13, 1, 26, 29],
 [15, 11, 26, 35],
 [14, 1, 26, 44],
 [14, 13, 26, 36],
 [14, 9, 26, 30],
 [13, 1, 26, 28],
 [13, 10, 26, 23],
 [14, 13, 26, 38],
 [15, 1, 26, 50],
 [15, 1, 26, 58],
 [15, 1, 26, 27],
 [14, 1, 26, 22],
 [13, 11, 26, 30],
 [15, 1, 26, 57],
 [13, 11, 26, 48],
 [14, 1, 26, 33],
 [15, 1, 26, 45],
 [14, 13, 26, 36],
 [14, 13, 26, 50],
 [13, 1, 26, 33],
 [15, 10, 26, 37],
 [12, 1, 26, 35],
 [13, 10, 26, 49],
 [14, 1, 26, 33],
 [13, 11, 26, 41],
 [15, 1, 26, 41],
 [13, 1, 26, 53],
 [14, 12, 26, 46],
 [12, 1, 26, 49],
 [14, 10, 26, 50],
 [14, 1, 26, 57],
 [14, 1, 26, 42],
 [12, 10, 26, 35],
 [14, 8, 26, 21],
 [11, 1, 26, 46],
 [12, 1, 26, 51],
 [13, 11, 26, 37],
 [13, 11, 26, 45],
 [14, 10, 26, 42],
 [14, 12, 26, 36],
 [13, 1, 26, 24],
 [15, 1, 26, 48],
 [11, 13, 26, 46],
 [14, 1, 26, 59],
 [13, 8, 26, 46],
 [13, 11, 26, 49],
 [14, 13, 26, 54],
 [14, 9, 26, 42],
 [11, 12, 26, 24],
 [14, 11, 26, 37],
 [15, 8, 26, 23],
 [14, 11, 26, 38],
 [14, 13, 26, 48],
 [14, 8, 26, 28],
 [14, 1, 26, 42],
 [14, 1, 26, 44],
 [13, 10, 26, 44],
 [13, 1, 26, 48],
 [13, 10, 26, 29],
 [14, 1, 26, 27],
 [13, 8, 26, 28],
 [14, 1, 26, 30],
 [12, 1, 26, 26],
 [13, 1, 26, 51],
 [15, 1, 26, 21],
 [12, 11, 26, 39],
 [14, 9, 26, 57],
 [13, 1, 26, 50],
 [14, 1, 26, 35],
 [14, 9, 26, 40],
 [14, 9, 26, 59],
 [13, 1, 26, 48],
 [14, 11, 26, 52],
 [14, 1, 26, 25],
 [14, 10, 26, 48],
 [14, 9, 26, 26],
 [14, 1, 26, 55],
 [12, 11, 26, 56],
 [13, 1, 26, 44],
 [15, 1, 26, 50],
 [13, 10, 26, 44],
 [13, 1, 26, 57],
 [15, 1, 26, 26],
 [14, 1, 26, 58],
 [14, 1, 26, 21],
 [12, 9, 26, 39],
 [14, 1, 26, 45],
 [13, 12, 26, 23],
 [13, 1, 26, 33],
 [12, 1, 26, 55],
 [15, 1, 26, 33],
 [13, 1, 26, 59],
 [11, 1, 26, 38],
 [15, 1, 26, 35],
 [13, 1, 26, 41],
 [12, 12, 26, 25],
 [14, 1, 26, 51],
 [13, 1, 26, 61],
 [12, 1, 26, 41],
 [15, 9, 26, 26],
 [14, 11, 26, 58],
 [13, 11, 26, 21],
 [14, 8, 26, 62],
 [14, 1, 26, 43],
 [15, 1, 26, 31],
 [12, 10, 26, 24],
 [12, 1, 26, 33],
 [13, 12, 26, 55],
 [15, 13, 26, 25],
 [13, 1, 26, 33],
 [12, 1, 26, 23],
 [14, 13, 26, 39],
 [11, 15, 26, 46],
 [12, 1, 26, 52],
 [14, 1, 26, 47],
 [14, 12, 26, 35],
 [14, 1, 26, 50],
 [14, 12, 26, 51],
 [11, 1, 26, 31],
 [14, 1, 26, 23],
 [13, 10, 26, 22],
 [13, 1, 26, 52],
 [13, 14, 26, 30],
 [14, 12, 26, 27],
 [15, 1, 26, 43],
 [14, 9, 26, 47],
 [14, 11, 26, 58],
 [14, 9, 26, 53],
 [13, 1, 26, 38],
 [15, 1, 26, 44],
 [15, 8, 26, 25],
 [13, 1, 26, 24],
 [13, 1, 26, 31],
 [14, 11, 26, 56],
 [14, 12, 26, 36],
 [15, 11, 26, 28],
 [14, 1, 26, 40],
 [13, 1, 26, 50],
 [15, 10, 26, 39],
 [13, 13, 26, 27],
 [14, 10, 26, 32],
 [14, 9, 26, 51],
 [14, 10, 26, 54],
 [12, 1, 26, 32],
 [14, 1, 26, 53],
 [12, 11, 26, 38],
 [12, 9, 26, 57],
 [14, 10, 26, 36],
 [13, 11, 26, 34],
 [14, 1, 26, 34],
 [13, 1, 26, 45],
 [13, 1, 26, 62],
 [14, 1, 26, 29],
 [10, 1, 26, 46],
 [13, 10, 26, 19],
 [12, 10, 26, 26],
 [12, 10, 26, 35]]
```

```python
your_length = [[len(logs()[i]['host']),len(logs()[i]['user_name']),len(logs()[i]['time']),
                len(logs()[i]['request'])] for i in range(len(logs()))] 

```

```python
for j in range(979):
    if solution_length[j]!=your_length[j]:
        print("Index number {}".format(j))
        print("Solution length {}".format(solution_length[j]))
        print("Your length {}".format(your_length[j]))
```

```python

```